---
title       : Logistic regression
description : Regression for binary outcomes, training and test sets and cross-validation

--- type:NormalExercise lang:r xp:100 skills:1 key:f8515d62e1
## Exercise 1

UNDER CONSTRUCTION, PLEASE DO NOT INTERACT HERE.

*** =instructions
- instruction

*** =hint
- hint

*** =pre_exercise_code
```{r}

```

*** =sample_code
```{r}
# Please do not interact here!
```

*** =solution
```{r}

```

*** =sct
```{r}

test_function("", args = "",incorrect_msg = "")
test_object("", incorrect_msg = "")
test_output_contains("", incorrect_msg = "")

test_error()
success_msg("Good work!")
```
